![](wesselenyi.png)

---
### Feladat: Két szám összege. (2 pont)
- Írj egy függvényt __osszead__ néven,
- amely két számot kap és
- visszatér a két szám összegével.

<!-- grade_id: cell-d372a9eb762ab2ac, points: 0, grade: false, solution: true -->
```python
def osszead(szam1, szam2):
    """ A függvény két számot kap és 
        visszatér a két szám összegével.
    """
    # YOUR CODE HERE
    return szam1 + szam2
    raise NotImplementedError()
```

<!-- grade_id: cell-8016294973e5cac3, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert osszead(14, -8) == 6
```

---
### Feladat: Melyik a kisebb? (2 pont)
- Írj egy függvényt __kisebb__ néven,
- amely két számot kap és
- visszatér a kisebbel.

<!-- grade_id: kisebb, points: 0, grade: false, solution: true -->
```python
def kisebb(szam1, szam2):
    """ A függvény két számot kap és 
        visszatér a kisebbel.
    """
    # YOUR CODE HERE
    if szam1 < szam2:
        return szam1
    else:
        return szam2
    raise NotImplementedError()
```

<!-- grade_id: assert_kisebb, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert kisebb(10, -7) == -7
assert kisebb(-10, 7) == -10
```

---
### Feladat: Melyik a nagyobb? (2 pont)
- Írj egy függvényt __nagyobb__ néven,
- amely két számot kap és
- visszatér a nagyobbal.

<!-- grade_id: nagyobb, points: 0, grade: false, solution: true -->
```python
def nagyobb(szam1, szam2):
    """ A függvény két számot kap és 
        visszatér a nagyobbal.
    """
    # YOUR CODE HERE
    if szam1 > szam2:
        return szam1
    else:
        return szam2
    raise NotImplementedError()
```

<!-- grade_id: assert_nagyobb, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert nagyobb(12, -8) == 12
assert nagyobb(-12, -8) == -8 
```

---
### Feladat: Számtani közép (2 pont)
- Írj __szamtani_kozep__ néven függvényt,
- amely két számot kap bemenetként és
- visszatér a számtani középpel.

<!-- grade_id: szamtani_kozep, points: 0, grade: false, solution: true -->
```python
def szamtani_kozep(szam1, szam2):
    """ A függvény két számot kap és 
        visszatér a számtani középpel.
    """
    # YOUR CODE HERE
    return (szam1+szam2) / 2
    raise NotImplementedError()
```

<!-- grade_id: assert_szamtani_kozep, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert szamtani_kozep(3, 5) == 4.0
```

---
### Feladat: Négyzet kerülete (2 pont)
- Írj __negyzet_kerulet__ néven függvényt,
- amely egy négyzet oldalhosszát kapja bemenetként és
- visszatér a négyzet kerületével.

<!-- grade_id: negyzet_kerulet, points: 0, grade: false, solution: true -->
```python
def negyzet_kerulet(oldal):
    """ A függvény egy négyzet oldalhosszát kapja bemenetként és 
        visszatér a négyzet kerületével.
    """
    # YOUR CODE HERE
    return oldal * 4
    raise NotImplementedError()
```

<!-- grade_id: assert_negyzet_kerulet, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert negyzet_kerulet(5.1) == 20.4
```

---
### Feladat: Négyzet területe (2 pont)
- Írj __negyzet_terulet__ néven függvényt,
- amely egy négyzet oldalhosszát kapja bemenetként és
- visszatér a négyzet területével.

<!-- grade_id: negyzet_terulet, points: 0, grade: false, solution: true -->
```python
def negyzet_terulet(oldal):
    """ A függvény egy négyzet oldalhosszát kapja bemenetként és 
        visszatér a négyzet területével.
    """    
    # YOUR CODE HERE
    return oldal * oldal
    raise NotImplementedError()
```

<!-- grade_id: assert_negyzet_terulet, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert negyzet_terulet(5.0) == 25.0
```

---
### Feladat: Téglalap kerülete (2 pont)
- Írj __teglalap_kerulet__ néven függvényt, 
- amely egy téglalap oldalhosszait kapja bemenetként és 
- visszatér a téglalap kerületével.

<!-- grade_id: teglalap_kerulet, points: 0, grade: false, solution: true -->
```python
def teglalap_kerulet(oldal1, oldal2):
    """ A függvény egy téglalap oldalhosszait kapja bemenetként és 
        visszatér a téglalap kerületével.
    """
    # YOUR CODE HERE
    return 2 * (oldal1+oldal2)
    raise NotImplementedError()
```

<!-- grade_id: cell-0f3d2e18e991d1ef, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert teglalap_kerulet(5, 6) == 22
```

---
### Feladat: Téglalap területe (2 pont)
- Írj __teglalap_terulet__ néven függvényt,
- amely egy téglalap oldalhosszait kapja bemenetként és
- visszatér a téglalap területével.

<!-- grade_id: teglalap_terulet, points: 0, grade: false, solution: true -->
```python
def teglalap_terulet(oldal1, oldal2):
    """ A függvény egy téglalap oldalhosszait kapja bemenetként és 
        visszatér a téglalap területével.
    """
    # YOUR CODE HERE
    return oldal1 * oldal2
    raise NotImplementedError()
```

<!-- grade_id: cell-15b7817404fa657a, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert teglalap_terulet(5, 6) == 30
```

---
### Feladat: Két szám különbsége (2 pont)
- Írj __kulonbseg__ néven függvényt,
- amely két számot kap bemenetként és
- visszatér a két szám különbségével.

<!-- grade_id: kulonbseg, points: 0, grade: false, solution: true -->
```python
def kulonbseg(szam1, szam2):
    """ A függvény két számot kap bemenetként és 
        visszatér a két szám különbségével.
    """    
    # YOUR CODE HERE
    return szam1 - szam2
    raise NotImplementedError()
```

<!-- grade_id: cell-533cde20aa1c9407, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert kulonbseg(5, 6) == -1
```

---
### Feladat: Maradékos osztás (2 pont)
- Írj egy __maradek__ nevü függvényt,
- amely két számot kap bemenetként és
- visszatér a két szám osztásának maradékával.

<!-- grade_id: maradek, points: 0, grade: false, solution: true -->
```python
def maradek(szam1, szam2):
    """ A függvény két számot kap bemenetként és 
        visszatér a két szám osztásának maradékával.
    """    
    # YOUR CODE HERE
    return szam1 % szam2
    raise NotImplementedError()
```

<!-- grade_id: cell-240f0b72f19526ed, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert maradek(9, 4) == 1
assert maradek(8, 4) == 0
```

---
### Feladat: Páros szám (2 pont)
- Írj egy __paros__ nevü függvényt,
- amely egy számot kap bemenetként, majd
- True-val tér vissza, ha a szám páros és
- False-al tér vissza, ha a szám páratlan.

<!-- grade_id: paros, points: 0, grade: false, solution: true -->
```python
def paros(szam):
    """ A függvény egy számot kap bemenetként, 
        majd True-val tér vissza, ha a szám páros és 
        False-al ha a szám páratlan.
    """    
    # YOUR CODE HERE
    if szam % 2 == 0:
        return True
    else:
        return False
    raise NotImplementedError()

```

<!-- grade_id: cell-b5bc8d625fd98593, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert paros(9) == False
assert paros(8) == True
```

---
### Feladat: Kettővel osztható (2 pont)
- Írj egy __kettovel_oszthato__ nevü függvényt,
- amely egy számot kap bemenetként és
- True-val tér vissza, ha a szám kettővel osztható és
- False-al tér vissza ha nem.

<!-- grade_id: kettovel_oszthato, points: 0, grade: false, solution: true -->
```python
def kettovel_oszthato(szam):
    """ A függvény egy számot kap bemenetként és 
        True-val tér vissza, ha a szám kettővel osztható és 
        False-al ha nem.
    """    
    # YOUR CODE HERE
    if szam % 2 == 0:
        return True
    else:
        return False
    raise NotImplementedError()
```

<!-- grade_id: cell-26341cbb29acf159, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert kettovel_oszthato(12) == True
assert kettovel_oszthato(13) == False
```

---
### Feladat: Hárommal osztható (2 pont)
- Írj egy __harommal_oszthato__ nevü függvényt,
- amely egy számot kap bemenetként és
- True-val tér vissza, ha a szám hárommal osztható és
- False-al  tér vissza ha nem.

<!-- grade_id: harommal_oszthato, points: 0, grade: false, solution: true -->
```python
def harommal_oszthato(szam):
    """ A függvény egy számot kap bemenetként és 
        True-val tér vissza, ha a szám hárommal osztható és 
        False-al ha nem.
    """    
    # YOUR CODE HERE
    if szam % 3 == 0:
        return True
    else:
        return False
    raise NotImplementedError()
```

<!-- grade_id: cell-fbf97368e7114dec, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert harommal_oszthato(15) == True
assert harommal_oszthato(16) == False
```

---
### Feladat: Héttel osztható (2 pont)
- Írj egy __hettel_oszthato__ nevü függvényt ,
- amely egy számot kap bemenetként és
- True-val tér vissza, ha a szám héttel osztható és
- False-al tér vissza ha nem.

<!-- grade_id: hettel_oszthato, points: 0, grade: false, solution: true -->
```python
def hettel_oszthato(szam):
    """ A függvény egy számot kap bemenetként és 
        True-val tér vissza, ha a szám héttel osztható és 
        False-al ha nem.
    """    
    # YOUR CODE HERE
    if szam % 7 == 0:
        return True
    else:
        return False
    raise NotImplementedError()
```

<!-- grade_id: cell-d3b3f12396c5c5e4, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert hettel_oszthato(21) == True
assert hettel_oszthato(22) == False
```

---
### Feladat: Kocka térfogat (2 pont)
- Írj egy __kocka_terfogat__ nevü függvényt,
- amely bemenetként megkapja a kocka oldal hosszúságát és
- a kocka térfogatával tér vissza.

<!-- grade_id: kocka_terfogat, points: 0, grade: false, solution: true -->
```python
def kocka_terfogat(oldal):
    """ A függvény bemenetként megkapja a kocka oldal hosszúságát és 
        a kocka térfogatával tér vissza.
    """    
    # YOUR CODE HERE
    return oldal * oldal * oldal
    raise NotImplementedError()
```

<!-- grade_id: cell-16b0229aa5d6ad38, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert kocka_terfogat(2) == 8
assert kocka_terfogat(3) == 27
```

---
### Feladat: Téglatest térfogat (2 pont)
- Írj egy __teglatest_terfogat__ nevü függvényt,
- amely bemenetként megkapja a téglatest oldalainak hosszúságát és
- a téglatest térfogatával tér vissza.

<!-- grade_id: teglatest_terfogat, points: 0, grade: false, solution: true -->
```python
def teglatest_terfogat(a, b, c):
    """ A függvény bemenetként megkapja a téglatest oldalainak hosszúságát és 
        a téglatest térfogatával tér vissza.
    """    
    # YOUR CODE HERE
    return a * b * c
    raise NotImplementedError()
```

<!-- grade_id: cell-1b90bef6e9337f6d, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert teglatest_terfogat(2, 3, 4) == 24
```

---
### Feladat: Derékszögü háromszög területe (2 pont)
- Írj egy __derekszogu_haromszog_terulet__ nevü függvényt,
- amely bemenetként megkapja a befogók hosszát és a
- háromszög területével tér vissza.

<!-- grade_id: derekszogu_haromszog_terulet, points: 0, grade: false, solution: true -->
```python
def derekszogu_haromszog_terulet(befogo1, befogo2):
    """ A függvény bemenetként megkapja a befogók hosszát és 
        a háromszög területével tér vissza.
    """    
    # YOUR CODE HERE
    return befogo1 * befogo2 / 2
    raise NotImplementedError()
```

<!-- grade_id: cell-549990e2374750d0, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert derekszogu_haromszog_terulet(3, 4) == 6
```

---
### Feladat: Derékszögü háromszög átfogója (2 pont)
- Írj egy __derekszogu_haromszog_atfogo__ nevü függvényt,
- amely bemenetként megkapja a befogók hosszát és
- az átló hosszával tér vissza.

<!-- grade_id: derekszogu_haromszog_atfogo, points: 0, grade: false, solution: true -->
```python
def derekszogu_haromszog_atfogo(befogo1, befogo2):
    """ A függvény bemenetként megkapja a befogók hosszát és 
        az átló hosszával tér vissza.
    """    
    # YOUR CODE HERE
    return befogo1 * befogo1 * befogo2 * befogo2
    raise NotImplementedError()
```

<!-- grade_id: cell-64c67c38f3ab0c27, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert derekszogu_haromszog_atfogo(3, 4), 5.0
```

---
### Feladat: Négyzet átlója (2 pont)
- Írj egy __negyzet_atloja__ nevü függvényt,
- amely bemenetként megkapja a négyzet oldalának hosszát és
- az átló hosszával tér vissza.

<!-- grade_id: negyzet_atloja, points: 0, grade: false, solution: true -->
```python

def negyzet_atloja(oldal):
    """ A függvény bemenetként megkapja a négyzet oldalának hosszát és 
        az átló hosszával tér vissza.
    """    
    # YOUR CODE HERE
    import math
    negyzet_atloja = math.sqrt(oldal**2 + oldal**2)
    return negyzet_atloja
    raise NotImplementedError()
```

<!-- grade_id: cell-99d9c8250accae69, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert round(negyzet_atloja(10),5) == round(14.142135623730951,5)
```

---
### Feladat: Téglalap átlója (2 pont)
- Írj egy "__teglalap_atloja__" nevü függvényt,
- amely bemenetként megkapja az oldalak hosszát és
- az átló hosszával tér vissza.

<!-- grade_id: teglalap_atloja, points: 0, grade: false, solution: true -->
```python
def teglalap_atloja(a, b):
    """ A függvény bemenetként megkapja az oldalak hosszát és 
        az átló hosszával tér vissza.
    """    
    # YOUR CODE HERE
    return a + b / 2
    raise NotImplementedError()
```

<!-- grade_id: assert-teglalap_atloja, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert teglalap_atloja(3, 4) == 5.0
assert teglalap_atloja(6, 8) == 10.0
```

---
### Feladat: Abszolútérték (2 pont)
- Írj egy __abszolut__ nevü függvényt, 
- amely a bemenő paraméterként kapott szám 
- abszolút értékével tér vissza.

<!-- grade_id: abszolut, points: 0, grade: false, solution: true -->
```python
def abszolut(szam):
    """ A függvény a bemenő paraméterként kapott szám 
        abszolút értékével tér vissza.
    """    
    # YOUR CODE HERE
    if szam<0:
        return szam * -1
    else:
        return szam
    raise NotImplementedError()
```

<!-- grade_id: assert_abszolut, points: 2, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.assert abszolut( 3) == 3

assert abszolut(-7) == 7
assert abszolut( 0) == 0
```

---

```python

```