![](wesselenyi.png)

---
### Feladat: Első karakter (2 pont)
- Írj egy függvényt __elso_karakter__ néven, amely 
- visszatér az adott string első karakterével!

<!-- grade_id: cell-f1801f964d2ea24d, points: 0, grade: false, solution: true -->
```python
def elso_karakter(string):
    """ Visszatér a string első karakterével.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-07167482fe4c9a96, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert elso_karakter("Hello") == "H"
```

---
### Feladat: Utolsó karakter (2 pont)
- Írj egy függvényt __utolso_karakter__ néven  
- amely visszatér az adott string utolso karakterével!

<!-- grade_id: cell-f7811faeecd5318e, points: 0, grade: false, solution: true -->
```python
def utolso_karakter(string):
    """ Visszatér az adott string utolso karakterével.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-94963cf785dc357b, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert utolso_karakter("Hello") == "o"
```

---
### Feladat: Legkisebb szám a listában.  (2 pont)
[Programozási tétel: Minimum kiválasztás]
- Készíts függvényt __legkisebb__ néven,  
- amely visszatér egy számokat tartalmazó lista legkisebb számával.
- Figyelem! A feladat megoldása során nem használhatod a min függvényt!

<!-- grade_id: cell-5525d313561ca77a, points: 0, grade: false, solution: true -->
```python
def legkisebb(lista):
    """ Visszatér egy számokat tartalmazó lista legkisebb számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-6e860b99809fa033, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert legkisebb( [7, 4, 9, -4, -8, 3]) == -8
```

---
### Feladat: Legnagyobb szám a listában.  (2 pont)
[Programozási tétel: Maximum kiválasztás]
- Készíts függvényt __legnagyobb__ néven,  
- amely visszatér egy számokat tartalmazó lista legnagyobb számával.
- Figyelem! A feladat megoldása során nem használhatod a max függvényt!

<!-- grade_id: cell-2eb1815d4f1cd448, points: 0, grade: false, solution: true -->
```python
def legnagyobb(lista):
    """ Visszatér egy számokat tartalmazó lista legnagyobb számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-97380de94390ab2c, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert legnagyobb( [7, 4, 9, -4, -8, 3] ) == 9
```

---
### Feladat: Listaban levő számok összege.  (2 pont)
[Programozási tétel: Összegzés]
- Készíts függvényt __osszeg__ néven,  
- amely visszatér egy számokat tartalmazó lista számainak összegével.
- Figyelem! A feladat megoldása során nem használhatod a sum függvényt!

<!-- grade_id: cell-b42288b4e32b1164, points: 0, grade: false, solution: true -->
```python
def osszeg(lista):
    """ Visszatér egy számokat tartalmazó lista számainak összegével.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-536323f5d5592173, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert osszeg( [1, 2, 3, 4, 5, 6] ) == 21
```

---       
### Feladat: Listában levő számok szorzata. (2 pont)
- Készíts függvényt __szorzat__ néven,  
- amely visszatér egy számokat tartalmazó lista számainak szorzatával.

<!-- grade_id: cell-8356d19cbdb50ac4, points: 0, grade: false, solution: true -->
```python
def szorzat(lista):
    """ Visszatér egy számokat tartalmazó lista számainak szorzatával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-a04d2a30d26ff7f9, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert szorzat( [1, 2, 3, 4, 5] )  == 120
```

---    
### Feladat: Páros számok száma a listában. (2 pont) 
[Programozási tétel: Megszámolás]
- Készíts függvényt __parosok_szama__ néven,  
- amely visszatér egy számokat tartalmazó lista páros számainak számával.

<!-- grade_id: cell-5120ec55a21e67ef, points: 0, grade: false, solution: true -->
```python
def parosok_szama(lista):
    """ Visszatér egy számokat tartalmazó lista páros számainak számával
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-2a8051b82bd4a682, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert parosok_szama( [7, 4, 9, -4, -8, 3, 1] ) == 3
```

---
### Feladat: Páratlan számok száma a listában.  (2 pont)
[Programozási tétel: Megszámolás]
- Készíts függvényt __paratlanok_szama__ néven,  
- amely visszatér egy számokat tartalmazó lista páros számainak számával.

<!-- grade_id: cell-c09dc4804a4ac6bc, points: 0, grade: false, solution: true -->
```python
def paratlanok_szama(lista):
    """ Visszatér egy számokat tartalmazó lista páros számainak számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-2e28d2490c8c34b8, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert paratlanok_szama( [7, 4, 9, -4, -8, 3, 1]) == 4
```

---
### Feladat: Pozitív számok száma a listában.  (2 pont)
[Programozási tétel: Megszámolás]
- Készíts függvényt __pozitivok_szama__ néven,  
- amely visszatér egy számokat tartalmazó lista pozitív számainak számával.

<!-- grade_id: cell-f029d6cb13c0f960, points: 0, grade: false, solution: true -->
```python
def pozitivok_szama(lista):
    """ visszatér egy számokat tartalmazó lista pozitív számainak számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-e20dbc160b0081cc, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert pozitivok_szama( [-7, -4, 9, -4, -8, 3, 1, 0]) == 3
```

---
### Feladat: Negatív számok száma a listában.  (2 pont)
[Programozási tétel: Megszámolás]
- Készíts függvényt __negativok_szama__ néven,  
- amely visszatér egy számokat tartalmazó lista negativ számainak számával.
    

<!-- grade_id: cell-8be228e7211a4eb4, points: 0, grade: false, solution: true -->
```python
def negativok_szama(lista):
    """ Visszatér egy számokat tartalmazó lista negativ számainak számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-6c7505d5ace32ac1, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert negativok_szama( [-7, -4, 9, -4, -8, 3, 1, 0]) == 4
```

---
### Feladat: Benne van a szám a listában?  (2 pont)
[Programozási tétel: Eldöntés]
Készíts függvényt __benne_van_a_listaban__ néven, 
- amelynek első paramétere egy számokat tartalmazó lista, 
- a második paramétere egy szám.
- A visszatérési érték __True__, ha  a szám benne van a listában.
- A visszatérési érték __False__, ha  a szám nics benne a listában.

<!-- grade_id: cell-cdeb4f150ca62d6d, points: 0, grade: false, solution: true -->
```python
def benne_van_a_listaban(lista, szam):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-e9f57f7d0b5cea6a, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert benne_van_a_listaban([-7, -4, 9, -4, -8, 3, 1, 0], 2) == False
assert benne_van_a_listaban([-7, -4, 9, -4, -8, 3, 1, 0], 9) == True
```

---
### Feladat: Benne van a betü a stringben?  (2 pont)
[Programozási tétel: Eldöntés]
Készíts függvényt __benne_van_a_stringben__ néven,  
- amelynek első paramétere egy string, 
- a második paramétere egy betü.
- A visszatérési érték __True__, ha  a betü benne van a stringben.
- A visszatérési érték __False__, ha  a betü nics benne a stringben.

<!-- grade_id: cell-8b65e5950aa64ec1, points: 0, grade: false, solution: true -->
```python
def benne_van_a_stringben(string, betu):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-6fd270aac89d06c5, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.   

assert benne_van_a_stringben("abrakadabra", "x") == False
assert benne_van_a_stringben("abrakadabra", "d") == True
```

---
### Feladat: Hányadik a listában?  (2 pont)
[Programozási tétel: Kiválasztás] 
- A listában a szám garantáltan megtalálható, nem kell vizsgálni a meglétét.
- Készíts függvényt __hanyadik_a_listaban__ néven,  
- amelynek első paramétere egy számokat tartalmazó lista, 
- a második paramétere egy szám.
- A visszatérési érték a paraméterként megadott szám első előfordulási helye a listában.

<!-- grade_id: cell-56f7f91594b049d7, points: 0, grade: false, solution: true -->
```python
def hanyadik_a_listaban(lista, szam):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-e41d3667b4b5994c, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert hanyadik_a_listaban([-7, -4, 9, -4, -8, 3, 1, 0], -7) == 0
assert hanyadik_a_listaban([-7, -4, 9, -4, -8, 3, 1, 0],  9) == 2
```

---
### Feladat: Hányadik a stringben?  (2 pont)
[Programozási tétel: Kiválasztás]
           A betü garantáltan megtalálható a stringben, nem kell vizsgálnunk a meglétét.
Készíts függvényt __hanyadik_a_stringben__ néven, 
- amelynek első paramétere egy string, 
- a második paramétere egy betü.
- A visszatérési érték a paraméterként megadott betü első előfordulási helye a stringben.

<!-- grade_id: cell-8554dd5a5d10285e, points: 0, grade: false, solution: true -->
```python
def hanyadik_a_stringben(string, betu):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-a755238535d248a7, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert hanyadik_a_stringben("abrakadabra", "a") == 0
assert hanyadik_a_stringben("abrakadabra", "k") == 4
```

---
### Feladat: Benne van-e a listában és hanyadik helyen?  (2 pont)
[Programozási tétel: Keresés]
Készíts függvényt __kereses_a_listaban__ néven,  
- amelynek első paramétere egy számokat tartalmazó lista, 
- a második paramétere egy szám.
- Ha a szám benne van a listában, akkor a visszatérési érték a paraméterként megadott szám első előfordulási helye a listában.
- A visszatérési érték None, ha  a szám nics benne a listában.

<!-- grade_id: cell-63d814eb10833526, points: 0, grade: false, solution: true -->
```python
def kereses_a_listaban(lista, szam):
    # YOUR CODE HERE
    raise NotImplementedError()
        
```

<!-- grade_id: cell-d772abc6c88e20b9, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert kereses_a_listaban([-7, -4, 9, -4, -8, 3, 1, 0], -7) == 0
assert kereses_a_listaban([-7, -4, 9, -4, -8, 3, 1, 0],  9) == 2
assert kereses_a_listaban([-7, -4, 9, -4, -8, 3, 1, 0],  2) == None
```

---
### Feladat: Hányadik a stringben?  (2 pont)
[Programozási tétel: Keresés]
Készíts függvényt __kereses_a_stringben__ néven, 
- amelynek első paramétere egy string, 
- a második paramétere egy betü.
- A visszatérési érték a paraméterként megadott betü első előfordulási helye a stringben. 
- A visszatérési érték None, ha  a betü nics benne a stringben.

<!-- grade_id: cell-b4465577eae39a2e, points: 0, grade: false, solution: true -->
```python
def kereses_a_stringben(string, betu):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ed85c5f7be855735, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert kereses_a_stringben("abrakadabra", "a") == 0
assert kereses_a_stringben("abrakadabra", "k") == 4
assert kereses_a_stringben("abrakadabra", "s") == None
```

---    
### Feladat: Páros számok kiválogatása egy listából.  (2 pont)
[Programozási tétel: Kiválogatás]
- Készíts függvényt __parosok_kivalogatasa néven__,
- amely visszatér egy listával, 
- amely a paraméterként átadott számokat tartalmazó lista páros számait tartalmazza.

<!-- grade_id: cell-0c25dfa5e4fba483, points: 0, grade: false, solution: true -->
```python
def parosok_kivalogatasa(lista):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-56c4b6fa09dd2b99, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert parosok_kivalogatasa( [7, 4, 9, 4, 8, 3, 1, 12, 0] ) == [4, 4, 8, 12, 0]
```

---
### Feladat: Páratlan számok számok kiválogatása egy listából. (2 pont)
[Programozási tétel: Kiválogatás]
- Készíts függvényt __paratlanok_kivalogatasa néven__,
- amely visszatér egy listával, 
- amely a paraméterként átadott számokat tartalmazó lista páratlan számait tartalmazza.

<!-- grade_id: cell-f27896f7360542b6, points: 0, grade: false, solution: true -->
```python
def paratlanok_kivalogatasa(lista):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-86c0a5263abff191, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert paratlanok_kivalogatasa( [7, 4, 9, 4, 8, 3, 1, 12, 0] ) == [7, 9, 3, 1]
```

---    
### Feladat: Pozitív számok kiválogatása egy listából.  (2 pont)
 [Programozási tétel: Kiválogatás] 
- Készíts függvényt __pozitivok_kivalogatasa__ néven,
- amely visszatér egy listával, 
- amely a paraméterként átadott számokat tartalmazó lista pozitiv számait tartalmazza. 

<!-- grade_id: cell-11cdfc8abf61b131, points: 0, grade: false, solution: true -->
```python
def pozitivok_kivalogatasa(lista):
    """ 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-6f9bda5fa89e1b50, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert pozitivok_kivalogatasa( [7, 4, 9, -4, -8, 3, 1, 12, 0] ) == [7, 4, 9, 3, 1, 12]
```

---
### Feladat: Negatív számok számok kiválogatása egy listából.  (2 pont)
    [Programozási tétel: Kiválogatás]
- Készíts függvényt __negatívok_kivalogatasa__ néven,
- amely visszatér egy listával, 
- amely a paraméterként átadott számokat tartalmazó lista negatív számait tartalmazza.

<!-- grade_id: cell-dddc5a3873f39fff, points: 0, grade: false, solution: true -->
```python
def negativok_kivalogatasa(lista):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-1842274518266bcb, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert negativok_kivalogatasa( [7, 4, 9, -4, -8, 3, 1, 12, 0] ) == [-4, -8]
```

---
### Feladat: Két halmaz metszete.  (2 pont)
[Programozási tétel: Metszet]
- Készíts függvényt __halmazok_metszete__ néven, 
- amely két paraméterként átadott halmaz azonos elemeivel mint halmazzal tér vissza.

<!-- grade_id: cell-d85477bc16033ed9, points: 0, grade: false, solution: true -->
```python
def halmazok_metszete(halmaz1, halmaz2):
    """ 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-4bee6b5bcd403b6c, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert halmazok_metszete({1, 2, 3, 4, 5, 6, 7}, {4, 5, 6, 7, 8, 9}) == {4, 5, 6, 7}
```

---
### Feladat: Két lista metszete.  (2 pont)
[Programozási tétel: Metszet]
- Készíts függvényt __listak_metszete__ néven, 
- amely két paraméterként átadott lista azonos elemeivel mint listával tér vissza.

<!-- grade_id: cell-769cc0a258329af6, points: 0, grade: false, solution: true -->
```python
def listak_metszete(lista1, lista2):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-092a1d2ce9714f19, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert set(listak_metszete([1, 2, 3, 4, 5, 6, 7], [4, 5, 6, 7, 8, 9])) == set([4, 5, 6, 7])
```

---
### Feladat: Két halmaz uniója.  (2 pont)
[Programozási tétel: Unió]
- Készíts függvényt __halmazok_unioja__ néven, 
- amely két paraméterként átadott halmaz minden elemeivel mint halmazzal tér vissza.

<!-- grade_id: cell-b14044b8a32ac060, points: 0, grade: false, solution: true -->
```python
def halmazok_unioja(halmaz1, halmaz2):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-5ee7efec77bbf9f2, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert halmazok_unioja({1, 2, 3, 4, 5, 6, 7}, {4, 5, 6, 7, 8, 9}) == {1, 2, 3, 4, 5, 6, 7, 8, 9}
```

---
### Feladat: Listaban levő számok átlaga (számtani közepe).  (2 pont)
Készíts függvényt __lista_atlag__ néven,  
amely visszatér egy számokat tartalmazó lista számainak átlagával.

<!-- grade_id: cell-631da50bc2cacab2, points: 0, grade: false, solution: true -->
```python
def lista_atlag(lista):
    """ Visszatér egy számokat tartalmazó lista számainak átlagával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-f7e43b7e8fde8d58, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert lista_atlag([1, 2, 3, 4, 5, 6]) == 21/6
```

---
### Feladat: Faktoriális.  (2 pont)
Készíts függvényt __faktorialis__ néven,  
amely visszatér a paraméterként megkapott szám faktoriálisával.

<!-- grade_id: cell-d2441539d171edae, points: 0, grade: false, solution: true -->
```python
def faktorialis(n):
    """ visszatér a paraméterként megkapott szám faktoriálisával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-dec0805197e67552, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert  faktorialis(0) == 1
assert  faktorialis(1) == 1
assert  faktorialis(2) == 1 * 2
assert  faktorialis(5) == 1 * 2 * 3 * 4 * 5
```

```python

```