![](wesselenyi.png)

---
### Feladat: Karakterek száma a fájlban.
- Írj egy függvényt __karakterek_szama__ néven 
- amely paraméterként egy fájlnevet kap és 
- visszatér a fájlban levő karakterek számával. ('\n' karakterekkel együtt)

<!-- grade_id: cell-bb376e7d044330b8, points: 0, grade: false, solution: true -->
```python
def karakterek_szama(fname):
    """ Visszatér a fájlban levő karakterek számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-de6a0a7f43b3020d, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert karakterek_szama("lorem.txt") == 18047

```

---  
### Feladat: Szavak száma a fájlban.
- Írj egy függvényt __szavak_szama__ néven 
- amely paraméterként egy fájlnevet kap és 
- visszatér a fájlban levő szavak számával.

<!-- grade_id: cell-daa0ffeb0a720516, points: 0, grade: false, solution: true -->
```python
def szavak_szama(fname):
    """ Visszatér a fájlban levő szavak számával
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7f65a4e846fe73d2, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert szavak_szama("lorem.txt") == 2304
```

--- 
### Feladat: Sorok száma a fájlban. 
- A __sorok_szama__ függvény 
- visszatér a  fájlban levő sorok számával.   

<!-- grade_id: cell-67c37d5d9365c773, points: 0, grade: false, solution: true -->
```python
def sorok_szama(fname):
    """ Visszatér a  fájlban levő sorok számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5283a66b81b454fd, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert sorok_szama("lorem.txt") == 82
```

---
### Feladat: r betük száma a fájlban. 
- Az __r_betuk_szama__ függvény, 
- visszatér a  fájlban levő 'r' betük számával.

<!-- grade_id: cell-729dfd47dd6a219b, points: 0, grade: false, solution: true -->
```python
def r_betuk_szama(fname):
    """ Visszatér a  fájlban levő 'r' betük számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6dabef711482b83b, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert r_betuk_szama("lorem.txt") == 790 
```

---     
###  Feladat: lorem szavak száma a fájlban. 
- A __lorem_szavak_szama__ függvény 
- visszatér a  fájlban levő "lorem" szavak számával.

<!-- grade_id: cell-057879e74fdb7802, points: 0, grade: false, solution: true -->
```python
def lorem_szavak_szama(fname):
    """ visszatér a  fájlban levő "lorem" szavak számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-45942e90c9d29bf6, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert lorem_szavak_szama("lorem.txt") == 27 
```

---   
### Feladat: A leggyakoribb karakter a fájlban. 
- A __leggyakoribb_karakter__ függvény 
- visszatér a  fájlban leggyakrabban előforduló karakterrel.

<!-- grade_id: cell-6d5ce7faa4093d7a, points: 0, grade: false, solution: true -->
```python
def leggyakoribb_karakter(fname):
    """ Visszatér a  fájlban leggyakrabban előforduló karakterrel.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0953654025282246, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert leggyakoribb_karakter("lorem.txt") ==  "i"
```

---
### Feladat: A leghosszabb sor hossza a fájlban. 
- A __leghosszabb_sor_hossza__ függvény 
- visszatér a  fájlban levő leghosszabb sor hosszával.

<!-- grade_id: cell-fe2563e13e08c7be, points: 0, grade: false, solution: true -->
```python
def leghosszabb_sor_hossza(fname):
    """ Visszatér a fájlban levő leghosszabb sor hosszával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1da6b9ac65b9b10f, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert leghosszabb_sor_hossza("lorem.txt") == 304
```

---
### Feladat: Téglalap osztály definiálása. [Objektumorientált programozás]
- Hozz létre egy osztályt __Teglalap__ néven.
- A Teglalap osztály lehetővé teszi a téglalap oldalhosszúságainak tárolását.
- A Teglalap osztály rendelkezik egy kerulet() nevü metódussal, amely az osztály segítségével létrehozott objektum metódusaként visszaadja az adott objektum kerületét.
- A Teglalap osztály rendelkezik egy terulet() nevü metódussal, amely az osztály segítségével létrehozott objektum metódusaként visszaadja az adott objektum területét.

<!-- grade_id: cell-6470cc86f346266c, points: 0, grade: false, solution: true -->
```python
class Teglalap:   
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-db6e05484c5bb48f, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.
assert Teglalap(3, 4).oldal1 == 3
assert Teglalap(3, 4).oldal2 == 4
assert Teglalap(3, 4).kerulet() == 14
assert Teglalap(3, 4).terulet() == 12
```

---
### Feladat: Négyzet osztály definiálása. 
[Objektumorientált programozás]
- Hozz létre egy osztályt __Negyzet__ néven.
- A Negyzet osztály lehetővé teszi a negyzet oldalhosszúságának tárolását.
- A Negyzet osztály rendelkezik egy kerulet() nevü metódussal, amely az osztály segítségével létrehozott objektum metódusaként visszaadja az adott objektum kerületét.
- A Negyzet osztály rendelkezik egy terulet() nevü metódussal, amely az osztály segítségével létrehozott objektum metódusaként visszaadja az adott objektum területét.

<!-- grade_id: cell-c50497c7afb2b424, points: 0, grade: false, solution: true -->
```python
class Negyzet:   
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b5c55e1b435109c0, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert Negyzet(3).kerulet() == 12
assert Negyzet(3).terulet() == 9
```

---
###  Feladat: Kocka osztály definiálása. 
[Objektumorientált programozás]
- Hozz létre egy osztályt __Kocka__ néven.
- A Kocka osztály lehetővé teszi a kocka oldalhosszúságának tárolását.
- A Kocka osztály rendelkezik egy __terfogat()__ nevü metódussal, amely az osztály segítségével létrehozott objektum metódusaként visszaadja az adott objektum térfogatát.
- A Kocka osztály rendelkezik egy __felszin()__ nevü metódussal, amely az osztály segítségével létrehozott objektum metódusaként visszaadja az adott objektum felszínét.

<!-- grade_id: cell-bd2597214ff09182, points: 0, grade: false, solution: true -->
```python
class Kocka:   
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-71fca5164aa5bb8e, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert Kocka(3).terfogat() == 27
assert Kocka(3).felszin() == 54
```

---
###  Feladat: String fájlba írása
- Készíts függvényt __string_fajlba__ néven,
- amely az első paraméterként kapott sztringet fájlba írja.
- A fájl nevét második paraméterként kapja meg a függvény.

<!-- grade_id: cell-d751cb7438289bc9, points: 0, grade: false, solution: true -->
```python
def string_fajlba(string, fajlnev):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-aea9af8a5da60e85, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert string_fajlba("csacska macska", "szoveg.txt") == None
assert open("szoveg.txt").read().strip() == "csacska macska"
```

---
###  Feladat: Számtani sorozat fájlba írása
- Készíts függvényt __szaz_szam_fajlba__ néven, 
- amely 1-tól 100-ig egyesével kiírja a számokat egy fájlba.
- Minden szám kerüljön új sorba.
- A fájl nevét paraméterként kapja meg a függvény.

<!-- grade_id: cell-add782eddf508282, points: 0, grade: false, solution: true -->
```python
def szaz_szam_fajlba(fajlnev):
    # YOUR CODE HERE
    raise NotImplementedError()


```

<!-- grade_id: cell-843285b165d8b7ef, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.assert szaz_szam_fajlba("szazas.txt") 

assert szaz_szam_fajlba("szazas.txt") == None
assert sum([int(i) for i in open("szazas.txt")]) == 5050
```

---
### Feladat: Első karakter a szövegfájlban
- Írj egy függvényt __elso_karakter_a_fajlban__ néven, 
- amely visszatér egy szövegfájl első karakterével.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-f96f206408697d8a, points: 0, grade: false, solution: true -->
```python
def elso_karakter_a_fajlban(fajlnev):
    """ Visszatér egy szövegfájl első karakterével.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d59255f150829afd, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert elso_karakter_a_fajlban("lorem.txt") == "L"
```

---
###  Feladat: Utolsó karakter a szövegfájlban
- Írj egy függvényt __utolso_karakter_a_fajlban__ néven, 
- amely visszatér egy szövegfájl utolsó karakterével.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-43f773e2052772e7, points: 0, grade: false, solution: true -->
```python
def utolso_karakter_a_fajlban(fajlnev):
    """ Visszatér egy szövegfájl utolsó karakterével
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02971e7540bad56b, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert utolso_karakter_a_fajlban("lorem.txt") == "."
```

---
###  Feladat: Számok összege egy szövegfájlban.
- Írj egy függvényt __szamok_osszege_a_fajlban__ néven 
- amely visszatér egy szövegfájlban levő számok összegével.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-49d5ebb5ad567c21, points: 0, grade: false, solution: true -->
```python
def szamok_osszege_a_fajlban(filename):
    """ Visszatér egy szövegfájlban levő számok összegével.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8fe0cd7b2ddc22a9, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert szamok_osszege_a_fajlban("szamok1.txt") == 16
```

---
###  Feladat: Számok átlaga egy szövegfájlban.
- Írj egy függvényt __szamok_atlaga_a_fajlban__ néven, 
- amely visszatér egy szövegfájlban levő számok átlagával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-6a6bf9ea1f725a10, points: 0, grade: false, solution: true -->
```python
def szamok_atlaga_a_fajlban(fajlnev):
    """ visszatér egy szövegfájlban levő számok átlagával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ed1c799f412dfb21, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert szamok_atlaga_a_fajlban("szamok1.txt") == 1.0
```

---
###  Feladat: Páros számok száma egy szövegfájlban.
- Írj egy függvényt __paros_szamok_szama_a_fajlban__ néven, 
- amely visszatér egy szövegfájlban levő páros számok számával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-ca3bd840842fbaf0, points: 0, grade: false, solution: true -->
```python
def paros_szamok_szama_a_fajlban(fajlnev):
    """ Visszatér egy szövegfájlban levő páros számok számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6ce134281a47781c, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.
        
assert paros_szamok_szama_a_fajlban("szamok1.txt") == 10
```

---
###  Feladat: Páratlan számok száma egy szövegfájlban.
- Írj egy függvényt __paratlan_szamok_szama_a_fajlban__ néven, 
- amely visszatér egy szövegfájlban levő páratlan számok számával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-7a3ad466ac8f60d3, points: 0, grade: false, solution: true -->
```python
def paratlan_szamok_szama_a_fajlban(fajlnev):
    """ 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b68c66ff9456144a, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert paratlan_szamok_szama_a_fajlban("szamok1.txt") == 6
```

---
###  Feladat: Pozitív számok száma egy szövegfájlban.
- Írj egy függvényt __pozitiv_szamok_szama_a_fajlban__ néven,
- amely visszatér egy szövegfájlban levő pozitiv számok számával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-19ffcaa4d9cfc7c3, points: 0, grade: false, solution: true -->
```python
def pozitiv_szamok_szama_a_fajlban(fajlnev):
    """  Visszatér egy szövegfájlban levő pozitiv számok számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7a2dc69320523bd1, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert pozitiv_szamok_szama_a_fajlban("szamok1.txt") == 10
```

---
###  Feladat: Negatív számok száma egy szövegfájlban.
- Írj egy függvényt __negativ_szamok_szama_a_fajlban__ néven, 
- amely visszatér egy szövegfájlban levő negativ számok számával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-8f85d488c68075bf, points: 0, grade: false, solution: true -->
```python
def negativ_szamok_szama_a_fajlban(fajlnev):
    """ Visszatér egy szövegfájlban levő negativ számok számával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fbc854a51eccbcf7, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert negativ_szamok_szama_a_fajlban("szamok1.txt") == 4
```

---
###  Feladat: Legkisebb szám egy szövegfájlban.
- Írj egy függvényt __legkisebb_szam_a_fajlban__ néven, 
- amely visszatér egy szövegfájlban levő lekisebb számmal.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-5c88b56c1822a52e, points: 0, grade: false, solution: true -->
```python
def legkisebb_szam_a_fajlban(fajlnev):
    """ Visszatér egy szövegfájlban levő lekisebb számmal.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ad1140ab4e1e4d0b, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert legkisebb_szam_a_fajlban("szamok1.txt") == -6
```

---
###  Feladat: Legnagyobb szám egy szövegfájlban.
- Írj egy függvényt __legnagyobb_szam_a_fajlban____ néven, 
- amely visszatér egy szövegfájlban levő legnagyobb számmal.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-7f7075c60b68973e, points: 0, grade: false, solution: true -->
```python
def legnagyobb_szam_a_fajlban(fajlnev):
    """ amely visszatér egy szövegfájlban levő 
        legnagyobb számmal.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-593cb9f48899c3d5, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert legnagyobb_szam_a_fajlban("szamok1.txt") == 4
```

---
###  Feladat: Párosok egy szövegfájlból.
- Írj egy függvényt __parosok_a_fajlbol__ néven, 
- amely visszatér a szövegfájlban levő páros számokkal mint listával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-786a1fae9a9ff9d1, points: 0, grade: false, solution: true -->
```python
def parosok_a_fajlbol(fajlnev):
    """ visszatér a szövegfájlban levő 
        páros számokkal mint listával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-287145df49ccf30e, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert parosok_a_fajlbol("szamok1.txt") == [4, 2, 4, 2, 0, -4, -6, 0, 4, 4]
```

---
###  Feladat: Páratlanok egy szövegfájlból.
- Írj egy függvényt __paratlanok_a_fajlbol__ néven, 
- amely visszatér a szövegfájlban levő páratlan számokkal mint listával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-a5d4b40561f97845, points: 0, grade: false, solution: true -->
```python
def paratlanok_a_fajlbol(fajlnev):
    """ visszatér a szövegfájlban levő 
        páratlan számokkal mint listával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8c5cae6bc0a09235, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert paratlanok_a_fajlbol("szamok1.txt") == [3, 1, 3, 1, -1, -1]
```

---
###  Feladat: Pozitívok egy szövegfájlból.
- Írj egy függvényt __pozitiv_a_fajlbol__ néven, 
- amely visszatér a szövegfájlban levő pozitiv számokkal mint listával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-244d3ffd521caa12, points: 0, grade: false, solution: true -->
```python
def pozitiv_a_fajlbol(fajlnev):
    """ visszatér a szövegfájlban levő 
        pozitiv számokkal mint listával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-dcfb6d8757e960bb, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert pozitiv_a_fajlbol("szamok1.txt") == [4, 3, 2, 1, 4, 3, 2, 1, 4, 4]
```

---
###  Feladat: Negatívok egy szövegfájlból.
- Írj egy függvényt __negativok_a_fajlbol__ néven, 
- amely visszatér a szövegfájlban levő negativ számokkal mint listával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-096a9e44642000d5, points: 0, grade: false, solution: true -->
```python
def negativok_a_fajlbol(fajlnev):
    """ visszatér a szövegfájlban levő 
        negativ számokkal mint listával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-dd8f6807bfe3c189, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert negativok_a_fajlbol("szamok1.txt") == [-1, -1, -4, -6] 
```

---
###  Feladat: Leggyakoribb szám a szövegfájlban.
- Írj egy függvényt __leggyakoribb_szam_a_fajlban__ néven, 
- amely visszatér a szövegfájlban levő leggyakoribb számmal.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-e3ee2d11618b6dc8, points: 0, grade: false, solution: true -->
```python
def leggyakoribb_szam_a_fajlban(fajlnev):
    """ visszatér a szövegfájlban levő leggyakoribb számmal.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-386b26477134a90a, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert leggyakoribb_szam_a_fajlban("szamok1.txt") == 4
```

---
###  Feladat: Hárommal osztható számok a szövegfájlban.
- Írj egy függvényt __harommal_oszthato_szamok_a_fajlban__ néven,
- amely visszatér a szövegfájlban levő hárommal osztható számok listájával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-7ffc5baa468ddacd, points: 0, grade: false, solution: true -->
```python
def harommal_oszthato_szamok_a_fajlban(fajlnev):
    """ visszatér a szövegfájlban levő 
        hárommal osztható számok listájával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-439c544493266350, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert harommal_oszthato_szamok_a_fajlban("szamok1.txt") == [3, 3, 0, -6, 0]
```

---
### Feladat: Neggyel osztható számok a szövegfájlban.

- Írj egy függvényt __neggyel_oszthato_szamok_a_fajlban__ néven, - amely visszatér a szövegfájlban levő neggyel osztható számok listájával.
- A függvény bemenő paramétere a fájl neve.

<!-- grade_id: cell-9b7eca8fbfdf74df, points: 0, grade: false, solution: true -->
```python
def neggyel_oszthato_szamok_a_fajlban(fajlnev):
    """ visszatér a szövegfájlban levő 
        neggyel osztható számok listájával.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9d090b4fd3dd29bc, points: 0, grade: true, solution: false -->
```python
# Ellenőrizd a függvényt a cella futtatásával! 
# Az assert csak hiba esetén ad visszajelzést.

assert neggyel_oszthato_szamok_a_fajlban("szamok1.txt") == [4, 4, 0, -4, 0, 4, 4]
```

```python

```